## Εργασία 1 ##

Καλωσήρθατε στην πρώτη σας εργασία. Η εργασία αυτή έχει σκοπό να σας βοηθήσει να εμπεδώσετε τα γραμμικά μοντέλα.

Τα Jupyter Notebooks είναι αλληλεπιδραστικά περιβάλλοντα προγραμματισμού ενσωματωμένα σε μια ιστοσελίδα. Στο μάθημα "Μηχανική Μάθηση" τα χρησιμοποιούμε στο πρακτικό μέρος του μαθήματος.

Στην εργασία αυτή θα πρέπει να συμπληρώσετε κώδικα Python 3 στα σημεία που αναφέρουν # YOUR CODE HERE. Μην τροποποιείτε τον κώδικα που βρίσκεται εκτός αυτών των περιοχών.

Πρωτού παραδόσετε την εργασία σας σιγουρευτείτε ότι ο κώδικας σε όλα τα κελιά τρέχει σωστά. Για το σκοπό αυτό επιλέξτε από το μενού **Χρόνος εκτέλεσης (runtime) -> Επανεκίνηση περιόδου λειτουργίας και εκτέλεση όλων**.

Συμπληρώστε το όνομα (**NAME**) και το **AEM** σας παρακάτω:

```python
NAME = "Ιωάννης Δεϊρμεντζόγλου"
AEM = "10015"
```

---
## Μέρος A: NumPy (1 μονάδα)

Η βιβλιοθήκη [**NumPy**](https://numpy.org/) είναι η βασική βιβιλιοθήκη για επιστημονικούς υπολογισμούς σε Python. Παρέχει ένα υψηλής επίδοσης αντικείμενο για την αποθήκευση πολυδιάστατων πινάκων και εργαλεία για επεξεργασίας αυτών των πινάκων.

Μπορείτε να συμβουλεύεστε την τεκμηρίωση ([NumPy Doc](https://numpy.org/doc/1.18/user/quickstart.html)) αυτών των βιβλιοθηκών κατά την εκπόνηση της εργασίας.

Κατά σύμβαση πάντα φορτώνουμε τη βιβλιοθήκη με τον ακόλουθο τρόπο:

```python
# Run this cell
import numpy as np
```

**A1**: Ορίστε συνάρτηση που επιστρέφει πίνακα 2 διδαστάσεων με r γραμμές και c στήλες και περιέχει τυχαίες τιμές από 0 έως 1 με τη βοήθεια της συνάρτησης np.random.rand. Ακολουθούν τεστ ορθής λειτουργίας της συνάρτησης.

<!-- grade_id: cell-4dd473645ec1eb64, points: 0, grade: false, solution: true -->
```python
def random_2d_array(r, c):
    """Επιστρέφει διδιάστατο πίνακα με r γραμμές και c στήλες.
    """
    return np.random.rand(r, c)
```

<!-- grade_id: random_2d_array_correct_shape, points: 5, grade: true, solution: false -->
```python
"""Τεστ επιστροφής πίνακα με τις σωστές διαστάσεις"""
assert np.shape(random_2d_array(3, 5)) == (3, 5)
assert np.shape(random_2d_array(5, 2)) == (5, 2)
```

**A2**: Ορίστε συνάρτηση που δέχεται δύο διδιάστατους πίνακες και επιστρέφει το γινόμενο του ανάστροφου του 1ου πίνακα με τον 2ο πίνακα με τη χρήση της numpy. Ακολουθούν τεστ ορθής λειτουργίας της συνάρτησης.

<!-- grade_id: cell-c697582f03491bb7, points: 0, grade: false, solution: true -->
```python
def multiply_transpose(a, b):
    """Επιστρέφει το γινόμενο του ανάστροφου του 1ου πίνακα με τον 2ο πίνακα.

    """
    return np.dot(a.T, b)
```

<!-- grade_id: cell-23e086f3e85d8f53, points: 5, grade: true, solution: false -->
```python
"""Τεστ επιστροφής πίνακα με τις σωστές διαστάσεις"""
a = random_2d_array(3, 5)
b = random_2d_array(3, 2)
assert np.shape(multiply_transpose(a, b)) == (5, 2)
assert np.shape(multiply_transpose(b, a)) == (2, 5)
```

## Μέρος Β: Γραμμικά μοντέλα (6 μονάδες) ##

Σε αυτό το μέρος της άσκησης θα χτίσουμε ένα μοντέλο γραμμικής παλινδρόμησης από το μηδέν.

**B1:** Υλοποιήστε τη σιγμοειδή συνάρτηση $\sigma(z)= \dfrac{1}{1 + exp(-z)}$ με τη NumPy. Ακολουθούν τεστ ορθής λειτουργίας της συνάρτησης.

<!-- grade_id: cell-c21b32f8218a444e, points: 0, grade: false, solution: true -->
```python
def sigmoid(z):
    """Υπολογίζει τη σιγμοειδή συνάρτηση."""
    return 1 / (1 + np.exp(-z))
```

<!-- grade_id: cell-c522a25040bf7202, points: 10, grade: true, solution: false -->
```python
"""Τεστ ορθού υπολογισμού σιγμοειδούς."""
assert sigmoid(0) == 0.5
assert np.round(sigmoid(7),8) == 0.99908895
assert np.round(sigmoid(-1),8) == 0.26894142
```

**B2:** Υλοποιήστε κλάση για τον αλγόριθμο της λογιστικής παλινδρόμησης. Αρχικοποιήστε τις παραμέτρους σε μηδενικές τιμές. Αποθηκεύστε τες στο διάνυσμα στήλη theta. Η συνάρτηση fit πρέπει να δέχεται πίνακα X με τις τιμές των μεταβλητών εισόδου και διάνυσμα στήλη y με τις τιμές της μεταβλητής εξόδου. Η συνάρτηση predict πρέπει να δέχεται πίνακα Χ με τις τιμές των μεταβλητών εισόδου. Ακολουθούν τεστ ορθής λειτουργίας των συναρτήσεων.

<!-- grade_id: cell-9df008253a524001, points: 0, grade: false, solution: true -->
```python
class MyLogisticRegression:
    def __init__(self, num_iterations=2000, learning_rate=0.004):
        self.theta = None
        self.num_iterations = num_iterations
        self.learning_rate = learning_rate

    def fit(self, X, y):
        m, n = X.shape
        self.theta = np.zeros((n + 1, 1))  # Update shape to include bias term

        # Add bias term to input matrix
        X = np.hstack((np.ones((m, 1)), X))

        for _ in range(self.num_iterations):
            z = np.dot(X, self.theta)
            h = sigmoid(z)
            gradient = np.dot(X.T, (h - y)) / m
            self.theta -= self.learning_rate * gradient

    def predict(self, X):
        if self.theta is None:
            raise ValueError("Model not trained yet. Call fit function first.")

        m = X.shape[0]
        X = np.hstack((np.ones((m, 1)), X))  # Add bias term
        z = np.dot(X, self.theta)
        h = sigmoid(z)
        predictions = (h >= 0.5).astype(int)
        return predictions
```

<!-- grade_id: cell-af8b565ded602dc1, points: 25, grade: true, solution: false -->
```python
"""Τεστ ορθής λειτουργίας fit"""
from sklearn.datasets import make_blobs

# δημιουργία δύο ομάδων από τυχαία σημεία με κέντρο 0,0,0 και 3,3,3 και τυπική απόκλιση 1
centers = [[0.0, 0.0, 0.0], [3.0, 3.0, 3.0]]
X, y = make_blobs(n_samples=200, centers=centers,
                  random_state=47, shuffle=False)
y = np.array(y).reshape(-1,1)

lr = MyLogisticRegression()
lr.fit(X, y)

assert np.allclose(lr.theta, np.array([[-1.1772064],[0.5026043],[0.40811666],[0.55685573]]))
```

<!-- grade_id: cell-e66896ed86d1a7ce, points: 25, grade: true, solution: false -->
```python
"""Τεστ ορθής λειτουργίας predict"""
from sklearn.datasets import make_blobs

# δημιουργία δύο ομάδων από τυχαία σημεία με κέντρο 0,0,0 και 3,3,3 και τυπική απόκλιση 1
centers = [[0.0, 0.0, 0.0], [3.0, 3.0, 3.0]]
X, y = make_blobs(n_samples=200, centers=centers,
                  random_state=47, shuffle=False)
y = np.array(y).reshape(-1,1)

lr = MyLogisticRegression()
lr.fit(X, y)

assert np.allclose(lr.predict(X[[0,199],:]), np.array([[0],[1]]))
```

## Μέρος Γ: Ομαλοποίηση (3 μονάδες)

**Γ1:** Υλοποιήστε κλάση για τον αλγόριθμο της λογιστικής παλινδρόμησης που να υποστηρίζει και ομαλοποίηση L1 και L2. Αρχικοποιήστε τις παραμέτρους σε μηδενικές τιμές. Αποθηκεύστε τες στο διάνυσμα στήλη theta. Η συνάρτηση fit πρέπει να δέχεται πίνακα X με τις τιμές των μεταβλητών εισόδου και διάνυσμα στήλη y με τις τιμές της μεταβλητής εξόδου. Ακολουθούν τεστ ορθής λειτουργίας.

<!-- grade_id: cell-8f651a74aaad362c, points: 0, grade: false, solution: true -->
```python
class MyLogisticRegression:
    def __init__(self, num_iterations=2000, learning_rate=0.004, penalty="None", l=10):
        self.theta = None
        self.num_iterations = num_iterations
        self.learning_rate = learning_rate
        self.penalty = penalty
        self.l = l

    def fit(self, X, y):
        m, n = X.shape
        self.theta = np.zeros((n + 1, 1))  # Update shape to include bias term

        # Add bias term to input matrix
        X = np.hstack((np.ones((m, 1)), X))

        for i in range(self.num_iterations):
            z = np.dot(X, self.theta)
            h = sigmoid(z)
            gradient = np.dot(X.T, (h - y)) / m

            if self.penalty == 'L1':
                regularization = (self.l / m) * np.sign(self.theta)
            elif self.penalty == 'L2':
                regularization = (self.l / m) * self.theta
            else:
                regularization = 0

            self.theta -= self.learning_rate * (gradient + regularization)

    def predict(self, X):
        if self.theta is None:
            raise ValueError("Model not trained yet. Call fit function first.")

        m = X.shape[0]
        X = np.hstack((np.ones((m, 1)), X))  # Add bias term
        z = np.dot(X, self.theta)
        h = sigmoid(z)
        predictions = (h >= 0.5).astype(int)
        return predictions


```

<!-- grade_id: cell-7b4a1e422511b0dd, points: 15, grade: true, solution: false -->
```python
"""Τεστ ορθής λειτουργίας fit L2"""
from sklearn.datasets import make_blobs

# δημιουργία δύο ομάδων από τυχαία σημεία με κέντρο 0,0,0 και 3,3,3 και τυπική απόκλιση 1
centers = [[0.0, 0.0, 0.0], [3.0, 3.0, 3.0]]
X, y = make_blobs(n_samples=200, centers=centers,
                  random_state=47, shuffle=False)
y = np.array(y).reshape(-1,1)

lr = MyLogisticRegression(penalty="L2", l=10)
lr.fit(X, y)
assert np.allclose(lr.theta, np.array([[-0.97858178],[0.44671962],[0.371177],[0.49835088]]))
lr = MyLogisticRegression(penalty="L2", l=100)
lr.fit(X, y)
assert np.allclose(lr.theta, np.array([[-0.30963538],[0.25825153],[0.23885811],[0.29480489]]))
```

<!-- grade_id: cell-2e2aab6691a11043, points: 15, grade: true, solution: false -->
```python
"""Τεστ ορθής λειτουργίας fit L1"""
from sklearn.datasets import make_blobs

# δημιουργία δύο ομάδων από τυχαία σημεία με κέντρο 0,0,0 και 3,3,3 και τυπική απόκλιση 1
centers = [[0.0, 0.0, 0.0], [3.0, 3.0, 3.0]]
X, y = make_blobs(n_samples=200, centers=centers,
                  random_state=47, shuffle=False)
y = np.array(y).reshape(-1,1)

lr = MyLogisticRegression(penalty="L1", l=100)
lr.fit(X, y)
assert np.allclose(lr.theta, np.array([[-0.0021886],[0.03118559],[0.01540396],[0.17818003]]))
```