## Εργασία 2 ##

Καλωσήρθατε στην δεύτερη σας εργασία. Η εργασία αυτή έχει σκοπό να σας βοηθήσει να εμπεδώσετε τα δενδρικά μοντέλα και την αξιολόγηση μοντέλων.

Στην εργασία αυτή θα πρέπει να συμπληρώσετε κώδικα Python 3 στα σημεία που αναφέρουν # YOUR CODE HERE. Μην τροποποιείτε τον κώδικα που βρίσκεται εκτός αυτών των περιοχών.

Πρωτού παραδόσετε την εργασία σας σιγουρευτείτε ότι ο κώδικας σε όλα τα κελιά τρέχει σωστά. Για το σκοπό αυτό επιλέξτε από το μενού Χρόνος εκτέλεσης (runtime) -> Επανεκίνηση περιόδου λειτουργίας και εκτέλεση όλων.

Συμπληρώστε το όνομα (NAME) και το AEM σας παρακάτω:

```python
NAME = "Δεϊρμεντζόγλου Ιωάννης"
AEM = "10015"
```

**1** Διαβάστε το διαθέσιμο από την sklearn σύνολο δεδομένων diabetes που αφορά την πρόβλεψη της πορείας που θα έχει ένας ασθενής με διαβήτη ένα χρόνο μετά (εργασία παλινδρόμησης) και αποθηκεύστε το σε μια μεταβλητή με το όνομα diabetes. (0.5 μονάδες)

<!-- grade_id: cell-ac5d304e50a110a6, points: 0, grade: false, solution: true -->
```python
from sklearn.datasets import load_diabetes

# Load the diabetes dataset
diabetes = load_diabetes()

# Now, 'diabetes' contains the data and target variables
X = diabetes.data  # Features
y = diabetes.target  # Target variable
```

<!-- grade_id: cell-2d158133c851eb4e, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθής ανάγνωσης του συνόλου δεδομένων"""
assert diabetes.feature_names == ['age', 'sex', 'bmi', 'bp', 's1', 's2', 's3', 's4', 's5', 's6']
```

**2** Χωρίστε το σύνολο που διαβάσατε σε υποσύνολο ελέγχου 20% και υποσύνολο εκπαίδευσης 80% με την μέθοδο train_test_split και σπόρο τυχαιότητας 42. Aποθηκεύστε το σύνολο εκπαίδευσης σε μεταβλητές X_train, y_train και το σύνολο ελέγχου σε μεταβλητές X_test, y_test. (0.5 μονάδες)

<!-- grade_id: cell-a637aedb66ae6d89, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

# Split the dataset into training and testing subsets
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)

# Now, X_train and y_train contain the training data and labels respectively,
# while X_test and y_test contain the testing data and labels respectively.
```

<!-- grade_id: cell-300331f01f617c3f, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == 0.02736
assert round(X_test[0][8], 5) == 0.03243
```

**3** Χρησιμοποιώντας τα δεδομένα εκπαίδευσης, εκπαιδεύστε δύο δένδρα παλινδρόμησης με κριτήρια διαχωρισμού τα squared_eror και poisson ([δείτε ενότητα 1.10.7.2](https://scikit-learn.org/stable/modules/tree.html#tree)) αντίστοιχα, θέτοντας τον σπόρο τυχαιότητας στην τιμή 42 και το μέγιστο βάθος στην τιμή 3. Υπολογίστε τη μετρική r2 με βάση τις προβλέψεις τους στα δεδομένα ελέγχου, τις οποίες αποθηκεύστε στις μεταβλητές y_pred1 (squared_error) και y_pred2 (poisson) αντίστοιχα.
Αποθηκεύστε τις μετρικές στις μεταβλητές r2_squared_error και r2_poisson αντίστοιχα. (1 μονάδα)

<!-- grade_id: cell-62cfb517c0f9271e, points: 0, grade: false, solution: true -->
```python
from sklearn.tree import DecisionTreeRegressor
from sklearn.metrics import r2_score

# Train regression tree with squared error separation criterion
regression_tree_squared_error = DecisionTreeRegressor(criterion='squared_error', max_depth=3, random_state=42)
regression_tree_squared_error.fit(X_train, y_train)
y_pred1 = regression_tree_squared_error.predict(X_test)
r2_squared_error = r2_score(y_test, y_pred1)

# Train regression tree with Poisson separation criterion
regression_tree_poisson = DecisionTreeRegressor(criterion='poisson', max_depth=3, random_state=42)
regression_tree_poisson.fit(X_train, y_train)
y_pred2 = regression_tree_poisson.predict(X_test)
r2_poisson = r2_score(y_test, y_pred2)
```

<!-- grade_id: cell-714ce0b2b256307e, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθού υπολογισμού των μετρικών"""
assert round(r2_squared_error, 3) == 0.329
assert round(r2_poisson, 3) == 0.391
assert len(y_pred1) == len(y_pred2)
```

**4** Αναπτύξτε κώδικα με βάση την ιδιότητα [tree_](https://scikit-learn.org/stable/auto_examples/tree/plot_unveil_tree_structure.html#sphx-glr-auto-examples-tree-plot-unveil-tree-structure-py) του παρακάτω μοντέλου, προκειμένου να εμφανίσετε το μονοπάτι που ακολουθήθηκε στο δένδρο για την πρόβλεψη της πρώτης περίπτωσης του συνόλου ελέγχου. Σε περίπτωση που ελέγχεται δεύτερη φορά μια μεταβλητή στο μονοπάτι, θα πρέπει να διατηρούμε τον τελευταίο και στενότερο περιορισμό. Ακολουθούν παραδείγματα ορθής λειτουργίας, στα οποία φαίνεται και η μορφή αν-τότε που πρέπει να έχει το μονοπάτι. (4 μονάδες)

```python
from sklearn.tree import export_text
model = DecisionTreeRegressor(max_depth=3, random_state=42)
model.fit(X_train, y_train)
```

Output:
```
DecisionTreeRegressor(max_depth=3, random_state=42)
```

<!-- grade_id: cell-1697542afe8f99a2, points: 0, grade: false, solution: true -->
```python
def get_path(tree, instance, feature_names):
    tree_ = tree.tree_
    node_indicator = tree.decision_path([instance])
    leaf_id = tree.apply([instance])[0]
    node_index = node_indicator.indices[node_indicator.indptr[0]: node_indicator.indptr[1]]

    conditions = {}

    for node_id in node_index:
        if leaf_id == node_id:
            break

        feature_index = tree_.feature[node_id]
        threshold = tree_.threshold[node_id]

        if instance[feature_index] <= threshold:
            symbol = "<="
            is_less_than = True
        else:
            symbol = ">"
            is_less_than = False

        # Update conditions dictionary with the narrowest constraint for each feature
        if feature_index not in conditions:
            conditions[feature_index] = (feature_index, symbol, threshold, is_less_than)
        else:
            _, _, existing_threshold, existing_is_less_than = conditions[feature_index]
            if (is_less_than and threshold < existing_threshold) or ( (not is_less_than) and threshold > existing_threshold):
                conditions[feature_index] = (feature_index, symbol, threshold, is_less_than)
            else :
             conditions[feature_index+len(node_index)] = (feature_index, symbol, threshold, is_less_than)
    path = ""
    for i, (feature_index, symbol, threshold, _) in enumerate(conditions.values()):
        if i > 0:
            path += " και "

        path += f"{feature_names[feature_index]} {symbol} {threshold:.2f}"

    path += f" τότε {round(tree_.value[leaf_id][0][0],2)}"

    # Adding the 'αν' (if) condition at the beginning of the path
    path = f"αν {path}"

    return path
```

<!-- grade_id: cell-8aabb13a60a743dd, points: 4, grade: true, solution: false -->
```python
"""Τεστ ορθού υπολογισμού του μονοπατιού"""
assert get_path(model, X_test[0], diabetes.feature_names) == 'αν bmi <= 0.01 και s5 > 0.01 και s4 <= 0.09 τότε 159.57'
assert get_path(model, X_test[1], diabetes.feature_names) == 'αν bmi > 0.01 και bmi <= 0.07 και s6 <= 0.03 τότε 175.8'
assert get_path(model, X_test[6], diabetes.feature_names) == 'αν bmi > 0.07 και s2 > 0.02 τότε 225.75'
```

**5** Διαχωρίστε τα δεδομένα diabetes σε υποσύνολα εκπαίδευσης 60% (X_train, y_train), επικύρωσης 15% (X_val, y_val) και ελέγχου 25% (X_test, y_test) με σπόρο τυχαιότητας 42. (1 μονάδα)

<!-- grade_id: cell-c2448c38ed5250e1, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

# First split to create training set and rest
X_train_rest, X_test, y_train_rest, y_test = train_test_split(diabetes.data, diabetes.target, test_size=0.25, random_state=42)

# Second split to create validation and test sets
X_train, X_val, y_train, y_val = train_test_split(X_train_rest, y_train_rest, test_size=0.2, random_state=42)
```

<!-- grade_id: cell-8da00aed9ea7c071, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == -0.08024
assert round(X_val[0][8], 5) == 0.09924
assert round(X_test[0][8], 5) == 0.03243
```

**6** Υλοποιήστε συνάρτηση, η οποία δοθέντος ενός δενδρικού μοντέλου παλινδρόμησης, διερευνά όλες τις ακέραιες τιμές της μεταβλητής min_samples_leaf από ένα κάτω όριο μέχρι ένα άνω όριο και επιστρέφει την βέλτιστη με βάση το mean squared error στο παραπάνω σύνολο επικύρωσης έπειτα από εκπαίδευση στο παραπάνω σύνολο εκπαίδευσης. Επιπλέον επιστρέφει και το mean squared error (στρογγυλοποιημένο σε ακέραιο) που πετυχαίνει στο σύνολο ελέγχου το μοντέλο με την βέλτιστη επίδοση στο σύνολο επικύρωσης. (2 μονάδες)

<!-- grade_id: cell-45e2ac4f9c5115f9, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import mean_squared_error

def best_min_samples_leaf(model, min, max):
    best_mse_val = float('inf')
    best_mse_test = float('inf')
    best_min_samples = -1


    for min_samples_leaf in range(min, max + 1):
        # Train model with current min_samples_leaf
        model.set_params(min_samples_leaf=min_samples_leaf)
        model.fit(X_train, y_train)

        # Evaluate model on validation set
        y_val_pred = model.predict(X_val)
        mse_val = mean_squared_error(y_val, y_val_pred)

        # Update best min_samples_leaf based on validation MSE
        if mse_val < best_mse_val:
            best_mse_val = mse_val
            best_min_samples = min_samples_leaf

            # Evaluate model on test set if validation MSE improves
            y_test_pred = model.predict(X_test)
            mse_test = mean_squared_error(y_test, y_test_pred)
            best_mse_test = mse_test

    return best_min_samples, int(round(best_mse_test))
```

<!-- grade_id: cell-0d6be21e54be9665, points: 2, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_depth=3, random_state=42)
assert best_min_samples_leaf(model, 1, 8) == (7, 3358)
assert best_min_samples_leaf(model, 5, 15) == (11, 3271)
```

**7** Υλοποιήστε συνάρτηση η οποία δοθέντος ενός μοντέλου παλινδρόμησης, ενός συνόλου δεδομένων και μίας μετρικής, το αξιολογεί με τη μέθοδο της σταυρωτής επικύρωσης "άφησε ένα εκτός". (1 μονάδα)

<!-- grade_id: cell-880c64392f85bccb, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import LeaveOneOut
import numpy as np

def leave_one_out(model, X, y, metric):
    loo = LeaveOneOut()
    scores = []

    for train_index, test_index in loo.split(X):
        X_train, X_test = X[train_index], X[test_index]
        y_train, y_test = y[train_index], y[test_index]

        model.fit(X_train, y_train)
        y_pred = model.predict(X_test)
        score = metric(y_test, y_pred)
        scores.append(score)

    return np.mean(scores)
```

<!-- grade_id: cell-4afd7b6de1cd1f7f, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_leaf_nodes=6, min_samples_leaf=2, random_state=42)
assert round(leave_one_out(model, diabetes.data, diabetes.target, mean_squared_error), 0) == 3354
```